```python
# Initialize Otter
import otter
grader = otter.Notebook("hw00.ipynb")
```

# Homework 0
## Intro to Jupyter Notebooks and Python
## Due Date: Tuesday, August 29th, 11:59 PM on Gradescope


### Detailed Submission Instructions Are Provided Provided at the end of this Notebook


## Assignment Scope:

1. Navigate Jupyter notebooks (like this one);
2. Write and evaluate some basic *expressions* in Python, the computer language of the course;
3. Call *functions* to use code other people have written; and
4. Complete the course intro survey

This hw covers parts of Chapter 3 of [this online textbook. ](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)


## Score Breakdown 
Question | Points
--- | ---
1a | ?
2a | 1
2b | 1
Total | ?




## Collaboration Policy

Data science is a collaborative activity. While you may talk with others about
the homework, we ask that you **write your solutions individually**. If you do
discuss the assignments with others please **include their names** at the top
of your notebook.

**Collaborators**: *list collaborators here*

## Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown



Question | Points
--- | ---
Question 0a | 1
Question 0b | 1
Question 1a | 2
Total | 40

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# 1. Jupyter Notebooks
This webpage is called a Jupyter Notebook. A notebook is a place to write programs and view their results, and also to write text.

### Jupyter Shortcuts ###

Here are some useful Jupyter notebook keyboard shortcuts.  To learn more keyboard shortcuts, go to **Help -> Keyboard Shortcuts** in the menu above. 

Here are a few we like:
1. `ctrl`+`return` : *Evaluate the current cell*
1. `shift`+`return`: *Evaluate the current cell and move to the next*
1. `esc` : *command mode* (may need to press before using any of the commands below)
1. `esc`+`a` : *create a cell above*
1. `esc`+`b` : *create a cell below*
1. `esc`+`dd` : *delete a cell*
1. `esc`+`m` : *convert a cell to markdown*
1. `esc`+`y` : *convert a cell to code*

## 1.1. Text cells
In a notebook, each rectangle containing text or code is called a *cell*.

Text cells (like this one) can be edited by double-clicking on them. They're written in a simple format called [Markdown](http://daringfireball.net/projects/markdown/syntax) to add formatting and section headings.  You don't need to learn Markdown, but you might want to.

After you edit a text cell, click the "run cell" button at the top that looks like â–¶| or hold down `shift` + `return` to confirm any changes. 

**Practice 1.1.1.** This paragraph is in its own text cell.  Try editing it so that this sentence is the last sentence in the paragraph, and then click the "run cell" â–¶| button or hold down `shift` + `return`.  This sentence, for example, should be deleted.  So should this one.

**Practice 1.1.2**:   Double click on this cell to see the markdown "behind the scenes"</span>:
 
This is type

<span style='color:Green'> This is GREEN type </span>

<span style='background :yellow' > This is highlighted type </span>

### This is big type
## This is bigger type
# This is even bigger type


*Italic*  or _Italic_ 

**Bold** 

***Bold and Italic*** 

~~Scratch Me.~~


## 1.2. Code cells
Other cells contain code in the Python 3 language. Running a code cell will execute all of the code it contains.

To run the code in a code cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle.  Next, either press â–¶| or hold down `shift` + `return`.

**Practice 1.2.1** Try running the cell below:

```python
print("Hello, World!")
```

And this one:

```python
print("\N{WAVING HAND SIGN}, \N{EARTH GLOBE ASIA-AUSTRALIA}!")
```

**Expressions:**  The fundamental building block of Python code is an expression. Cells can contain multiple lines with multiple expressions. When you run a cell, the lines of code are executed in the order in which they appear. Every `print` expression prints a line. 

Run the next cell and notice the order of the output.

```python
print("First this line is printed,")
print("and then this one.")
```

**Practice 1.2.2:**  Change the code in the cell above to print out the following:

    First this line,
    then the whole ðŸŒ,
    and then this one.


### Viewing Documentation

To output the documentation for a function, use the `help` function.

```python
help(print)
```

You can also use Jupyter to view function documentation inside your notebook. The function must already be defined in the kernel for this to work.

**Practice 1.2.3:** Below, click your mouse anywhere on the `print` block below and use `Shift` + `Tab` to view the function's documentation. 

```python
print('Welcome to CSCI 3022.')
```

## 1.3. Writing Jupyter notebooks
You can use Jupyter notebooks for your own projects or documents.  When you make your own notebook, you'll need to create your own cells for text and code.

To add a cell, click the + button in the menu bar.  It'll start out as a text cell.  You can change it to a code cell by clicking inside it so it's highlighted, clicking the drop-down box next to the restart (âŸ³) button in the menu bar, and choosing "Code".

**Practice 1.3.1.** Add a code cell below this one.  Write code in it that prints out:
   
    A whole new cell! â™ªðŸŒâ™ª

(That musical note symbol is like the Earth symbol.  Its long-form name is `\N{EIGHTH NOTE}`.)

Run your cell to verify that it works.

## 1.4. Errors
Python is a language, and like natural human languages, it has rules.  It differs from natural language in two important ways:
1. The rules are *simple*.  You can learn most of them in a few weeks and gain reasonable proficiency with the language in a semester.
2. The rules are *rigid*.  If you're proficient in a natural language, you can understand a non-proficient speaker, glossing over small mistakes.  A computer running Python code is not smart enough to do that.

Whenever you write code, you'll make mistakes.  When you run a code cell that has errors, Python will sometimes produce error messages to tell you what you did wrong.

Errors are okay; even experienced programmers make many errors.  When you make an error, you just have to find the source of the problem, fix it, and move on.

```python
print("This line is missing something."
```

**Note:** In the toolbar, there is the option to click `Cell > Run All`, which will run all the code cells in this notebook in order. However, the notebook stops running code cells if it hits an error, like the one in the cell above.

You should see something like this (minus our annotations):

<img src="error.jpg"/>

The last line of the error output attempts to tell you what went wrong.  The *syntax* of a language is its structure, and this `SyntaxError` tells you that you have created an illegal structure.  The message is saying Python expected you to write something more (in this case, a right parenthesis) before finishing the cell.

There's a lot of terminology in programming languages, but you don't need to know it all in order to program effectively. If you see a cryptic message like this, you can often get by without deciphering it.  (Of course, if you're frustrated, ask a neighbor or a staff member for help.) 


**Practice 1.4.1** Try to fix the code above so that you can run the cell and see the intended message instead of an error.

## 1.5. The Kernel
The kernel is a program that executes the code inside your notebook and outputs the results. In the top right of your window, you can see a circle that indicates the status of your kernel. If the circle is empty (âšª), the kernel is idle and ready to execute code. If the circle is filled in (âš«), the kernel is busy running some code. 

Next to every code cell, you'll see some text that says `[ ]`. Before you run the cell, you'll see `[ ]`. When the cell is running, you'll see `[*]`. If you see an asterisk (\*) next to a cell that doesn't go away, it's likely that the code inside the cell is taking too long to run, and it might be a good time to interrupt the kernel (discussed below). When a cell is finished running, you'll see a number inside the brackets, like so: `[1]`. The number corresponds to the order in which you run the cells; so, the first cell you run will show a 1 when it's finished running, the second will show a 2, and so on. 

You may run into problems where your kernel is stuck for an excessive amount of time, your notebook is very slow and unresponsive, or your kernel loses its connection. If this happens, try the following steps:
1. At the top of your screen, click **Kernel**, then **Interrupt**.
2. If that doesn't help, click **Kernel**, then **Restart**. If you do this, you will have to run your code cells from the start of your notebook up until where you paused your work.

# 2. Number Types

Quantitative information arises everywhere in data science. In addition to representing commands to print out lines, expressions can represent numbers and methods of combining numbers.


Python has two real number types:

`int`: 		an integer of any size

`float`:	a number with an optional fractional part

An int never has a decimal point; a float always does

A float might be printed using scientific notation




The expression `3.2500` evaluates to the floating point 3.25. (Run the cell and see.)

```python
3.2500
```

Notice that we didn't have to `print`. When you run a notebook cell, if the last line has a value, then Jupyter helpfully prints out that value for you. However, it won't print out prior lines automatically.

```python
print(2)
3
4
```

Above, you should see that 4 is the value of the last expression, 2 is printed, but 3 is lost forever because it was neither printed nor last.

You don't want to print everything all the time anyway.  But if you feel sorry for 3, change the cell above to print it.

## 2.1 Limitations of float values:

Be aware that float values have the following limitations, which can lead to issues if you're not aware:

- They have limited precision of 15-16 decimal places

- After arithmetic, the final few decimal places can be wrong

- They have limited size (but the limit is huge)


**Practice 2.2.1**:  Run the cells below to observe some of the issues that arise when working with float values

```python
382764238764.238742034876503865238723985 - 382764238764.238742
```

```python
print(0.1 + 0.2)
```

```python
0.1 + 0.2 == 0.3
```

# 3. Names and Assignment Statements
In natural language, we have terminology that lets us quickly reference very complicated concepts.  We don't say, "That's a large mammal with brown fur and sharp teeth!"  Instead, we just say, "Bear!"

In Python, we do this with *assignment statements*. An assignment statement has a name on the left side of an `=` sign and an expression to be evaluated on the right.

```python
ten = 3 * 2 + 4
```

When you run that cell, Python first computes the value of the expression on the right-hand side, `3 * 2 + 4`, which is the number 10.  Then it assigns that value to the name `ten`.  At that point, the code in the cell is done running.

After you run that cell, the value 10 is bound to the name `ten`:

```python
ten
```

The statement `ten = 3 * 2 + 4` is not asserting that `ten` is already equal to `3 * 2 + 4`, as we might expect by analogy with math notation.  Rather, that line of code changes what `ten` means; it now refers to the value 10, whereas before it meant nothing at all.

If the designers of Python had been ruthlessly pedantic, they might have made us write

    define the name ten to hereafter have the value of 3 * 2 + 4 

instead.  You will probably appreciate the brevity of "`=`"!  But keep in mind that this is the real meaning.

**Practice 3.1.** Run the following cell which uses a variable name `eleven` that hasn't been assigned to anything. You'll see an error!

```python
eleven + 8
```

A common pattern in Jupyter notebooks is to assign a value to a name and then immediately evaluate the name in the last line in the cell so that the value is displayed as output. 

```python
close_to_pi = 355/113
close_to_pi
```

Another common pattern is that a series of lines in a single cell will build up a complex computation in stages, naming the intermediate results.

```python
semimonthly_salary = 842.5
monthly_salary = 2 * semimonthly_salary
number_of_months_in_a_year = 12
yearly_salary = number_of_months_in_a_year * monthly_salary
yearly_salary
```

Names in Python can have letters (upper- and lower-case letters are both okay and count as different letters), underscores, and numbers.  The first character can't be a number (otherwise a name might look like a number).  And names can't contain spaces, since spaces are used to separate pieces of code from each other.

Other than those rules, what you name something doesn't matter *to Python*.  For example, this cell does the same thing as the above cell, except everything has a different name:

```python
a = 842.5
b = 2 * a
c = 12
d = c * b
d
```

**However**, names are very important for making your code *readable* to yourself and others.  The cell above is shorter, but it's totally useless without an explanation of what it does.

## 3.1. Checking Your Code

Now that you know how to name things, you can start using the built-in *tests* to check whether your work is correct. Sometimes, there are multiple tests for a single question, and passing all of them is required to receive credit for the question. **Some questions will also contain hidden tests that won't be run until you submit to Gradescope.** Please don't change the contents of the test cells. 

Go ahead and attempt Question 1 below. Running the cell directly after it will test whether you have assigned `seconds_in_a_decade` correctly in Question 1. If you haven't, this test will tell you the correct answer. Resist the urge to just copy it, and instead try to adjust your expression. (Sometimes the tests will give hints about what went wrong...)


### Question 1 (xx points)**
Assign the name `seconds_in_a_decade` to the number of seconds between midnight January 1, 2010 and midnight January 1, 2020. Note that there are two leap years in this span of a decade. A non-leap year has 365 days and a leap year has 366 days.

Replace the ellipses (`...`) with your expression.


```python
# Change the next line 
# so that it computes the number of seconds in a decade 
# and assigns that number the name, seconds_in_a_decade.

seconds_in_a_decade = ...

# We've put this line in this cell 
# so that it will print the value you've given to seconds_in_a_decade when you run it.  
# You don't need to change this.
seconds_in_a_decade
```

```python
grader.check("q1")
```

### Question 2 (xx pts)
When you run the following cell, Python produces a cryptic error message.

```python
4 = 2 + 2
```

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `names_q2` below to indicate your answer. **(4 Points)**

1. Python is smart and already knows `4 = 2 + 2`.

2. In Python, it's a rule that the `=` sign must have a variable name to its left, and `4` isn't a variable name.

3. It should be `2 + 2 = 4`.

4. I don't get an error message. This is a trick question.


```python
names_q2 = ...
```

```python
grader.check("q2")
```

## 2.1. Operators


Many basic arithmetic operations are built into Python.  The textbook section on [Expressions](http://www.inferentialthinking.com/chapters/03/1/Expressions.html) describes all some key arithmetic operators.  The common operator that differs from typical math notation is `**`, which raises one number to the power of the other. So, `2**3` stands for $2^3$ and evaluates to 8. 

**Practice:** Play around with the operators $\%, /, **$, and $//$ to see what they do:

```python

```

The order of operations is the same as what you learned in elementary school, and Python also has parentheses.  For example, compare the outputs of the cells below. The second cell uses parentheses for a happy new school year!

```python
7+6*5-6*3**2*2**3/4*7
```

```python
7+(6*5-(6*3))**2*((2**3)/4*7)
```

In standard math notation, the first expression is

$$6 + 6 \times 5 - 6 \times 3^2 \times \frac{2^3}{4} \times 7,$$

while the second expression is

$$6 + (6 \times 5 - (6 \times 3))^2 \times \left(\frac{\left(2^3\right)}{4} \times 7\right).$$

### Question 3 ###

Write a Python expression in this next cell that's equal to $\displaystyle 5 \times \left(3 \frac{10}{11}\right) - 50 \frac{1}{3} + 2^{0.5 \times 22} - \frac{7}{33} + 6$.  That's five times three and ten elevenths, minus fifty and a third, plus two to the power of half twenty-two, minus seven thirty-thirds plus six.  By "$\displaystyle 3 \frac{10}{11}$" we mean $\displaystyle 3+\frac{10}{11}$, not $\displaystyle 3 \times \frac{10}{11}$.

Replace the ellipses (`...`) with your expression.  Try to use parentheses only when necessary.

Go ahead and attempt Question 1. Running the cell directly after it will test whether you have assigned `q1_answer` correctly in Question 1. 


```python
q1_answer = ...
```

```python
grader.check("q3")
```

**Question 4:** Write a function `square` that returns the square of its argument.

```python
def square(x):
    ...
```

```python
grader.check("q4")
```

<!-- BEGIN QUESTION -->

**Question**: What is $\int e^{x/4} \,dx$?
Write your answer in the cell below using Markdown and LaTeX

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 0!

### Submission Instructions

Below, you will see two cells.
* Running the first cell will automatically generate a PDF of your answers to all questions that need to be manually graded.
* Running the second cell will automatically generate a zip with your autograded answers.

**You are responsible for submitting **both** the coding and written portions of Homework 0 to their respective Gradescope portals:
* **Homework 0 Coding**: Submit your Jupyter notebook zip file for Homework 0,
which can be generated and downloaded from DataHub by using the `grader.export()`
cell provided below.
* **Homework 0 Written**: Submit your hw00.PDF to Gradescope 
which can be generated and downloaded from DataHub by using the first 
cell provided below.

To receive credit on this assignment, **you must submit both your coding and written
portions to their respective Gradescope portals.** Your written submission (a single
PDF) can be generated as follows:

1. Access your answers to manually graded Homework 1A  questions in one of three ways:
      * <b>*Automatically create PDF (recommended)*</b>: Run the first cell below and download the generated PDF. This function will extract your response to the manually graded questions and put them on separate page. This process may fail if your answer is not properly formatted; if this is the case, check out common errors and solution described on Ed or follow either of the two ways described below.
      * *Manually download PDF*: If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> PDF`. If you choose to go this route, you must take special care to ensure all appropriate pages are chosen for each question on Gradescope.
      * *Take screenshots*: If that doesn't work either, you can take screenshots of your answers (and your code if present) to manually graded questions and include them as images in a PDF. The manually graded questions are listed at the top of the Homework 1A notebook.
    

1. **Important**: When submitting on Gradescope, you **must tag pages to each question correctly** (it prompts you to do this after submitting your work). This significantly streamlines the grading process for our readers. Failure to do this may result in a score of 0 for untagged questions.

**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Before running anything, save your notebook.  AFTER running the cell below, click on <a href='hw00.pdf' download>this link to download the PDF to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope </a>.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 1)
```
>>> # It looks like you didn't give anything the name
>>> # seconds_in_a_decade.  Maybe there's a typo, or maybe you
>>> # just need to run the cell below Question 1 where you defined
>>> # seconds_in_a_decade.  (Click that cell and then click the "run
>>> # cell" button in the menu bar above.)
>>> assert 'seconds_in_a_decade' in vars()
```
```
>>> # It looks like you didn't change the cell to define
>>> # seconds_in_a_decade appropriately.  It should be a number,
>>> # computed using Python's arithmetic.  For example, this is
>>> # almost right:
>>> #   seconds_in_a_decade = 10*365*24*60*60
>>> assert seconds_in_a_decade != ...
```
```
>>> # It looks like you didn't account for leap years.
>>> # There were 2 leap years and 8 non-leap years in this period.
>>> # Leap years have 366 days instead of 365.
>>> assert seconds_in_a_decade != 315360000
```
```
>>> assert seconds_in_a_decade == 315532800
```


Tests for q2 (points: [0])
```
>>> assert 1 <= names_q2 <= 4
```


Tests for q3 (points: 1)
```
>>> assert q1_answer == 2023.0
```


Tests for q4 (points: 1)
```
>>> assert square(1) == 1
```
```
>>> assert square(0) == 0
```
